# Adaboost
For this exercise you will implement AdaBoost from scratch and applied it to a spam dataset. You will be classifying data into spam and not spam. You can call DecisionTreeClassifier from sklearn to learn your base classifiers.

Here is how you train a decision tree classifier with weights.

`
h = DecisionTreeClassifier(max_depth=1, random_state=0)
h.fit(X, Y, sample_weight=w)
`

```python
import numpy as np
from sklearn.tree import DecisionTreeClassifier
```

```python
# accuracy computation
def accuracy(y, pred):
    return np.sum(y == pred) / float(len(y)) 
```

<!-- grade_id: cell-1eaf818528c8d676, points: 0, grade: false, solution: true -->
```python
def parse_spambase_data(filename):
    """ Given a filename return X and Y numpy arrays

    X is of size number of rows x num_features
    Y is an array of size the number of rows
    Y is the first element of each row. (Convert 0 to -1)
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    return X, Y
```

<!-- grade_id: cell-73926460e1f70f54, points: 1, grade: true, solution: false -->
```python
y_test = np.array([1., -1., 1., 1., -1., -1., 1., 1., 1., -1.])
X, Y = parse_spambase_data("tiny.spam.train")
for i in range(len(y_test)): assert(y_test[i] == Y[i])
n, m = X.shape
assert(n == 10)
assert(m == 57)
```

<!-- grade_id: cell-befcece7be9c6839, points: 0, grade: false, solution: true -->
```python
def adaboost(X, y, num_iter):
    """Given an numpy matrix X, a array y and num_iter return trees and weights 
   
    Input: X, y, num_iter
    Outputs: array of trees from DecisionTreeClassifier
             trees_weights array of floats
    Assumes y is {-1, 1}
    """
    trees = []
    trees_weights = [] 
    N, _ = X.shape
    d = np.ones(N) / N

    # YOUR CODE HERE
    raise NotImplementedError()
    return trees, trees_weights
```

<!-- grade_id: cell-4dc8edbf0e11fab8, points: 1, grade: true, solution: false -->
```python
X, Y = parse_spambase_data("tiny.spam.train")
trees, weights = adaboost(X, Y, 2)
assert(len(trees) == 2)
assert(len(weights) == 2)
assert(isinstance(trees[0], DecisionTreeClassifier))
```

<!-- grade_id: cell-8702186fedbd58e1, points: 2, grade: true, solution: false -->
```python
x = np.array([[0, -1], [1, 0], [-1, 0]])
y = np.array([-1, 1, 1])
trees, weights = adaboost(x, y, 1)
h = trees[0]
pred = h.predict(x)
for i in range(len(y)): assert(pred[i] == y[i])
```

<!-- grade_id: cell-50ad4a5c81e7c016, points: 0, grade: false, solution: true -->
```python
def adaboost_predict(X, trees, trees_weights):
    """Given X, trees and weights predict Y
    """
    # X input, y output
    N, _ =  X.shape
    y = np.zeros(N)
    # YOUR CODE HERE
    raise NotImplementedError()
    return y
```

<!-- grade_id: cell-7f28f00061f7b5a5, points: 3, grade: true, solution: false -->
```python
x = np.array([[0, -1], [1, 0], [-1, 0]])
y = np.array([-1, 1, 1])
trees, weights = adaboost(x, y, 1)
pred = adaboost_predict(x, trees, weights)
for i in range(len(y)):
    assert(pred[i] == y[i])
```

<!-- grade_id: cell-89126a5a7f8f0e1b, points: 3, grade: true, solution: false -->
```python
X, Y = parse_spambase_data("spambase.train")
X_test, Y_test = parse_spambase_data("spambase.test")
trees, trees_weights = adaboost(X, Y, 10)
Yhat = adaboost_predict(X, trees, trees_weights)
Yhat_test = adaboost_predict(X_test, trees, trees_weights)
    
acc_test = accuracy(Y_test, Yhat_test)
acc_train = accuracy(Y, Yhat)
print("Train Accuracy %.4f" % acc_train)
print("Test Accuracy %.4f" % acc_test)
assert(np.around(acc_train, decimals=4)==0.9111)
assert(np.around(acc_test, decimals=4)==0.9190)
```